```python

```

You should have read all the readings and content from the lectures given on Brightspace.


In order to make it easy to grade your works, here are the notes you should read and follow carrefully.  We will take off points for not following them. We will grade based on assuming you have read and understood them.
    
1. **Using the Notebook to Show Your Work**: You must learn to write code in the notebook... It is a core tool for data science and will make it easier to develop and document your work if you become good at using it. Writing your code in another tool and pasting it into the notebook will probably not work well: Some of you forgot to include code elsewhere, or the spacing became messed up, or it didn't run as copied.  You must be sure your code cells execute because we will test them.  So learn how to run code in the notebook cells to double check your work.
2. **Read Directions Carefully**:  The instructions for your code are very important. If you don't follow the requirements, your application won't do as requested. Making it work correctly is part of learning to program.  If we worded something unclearly, send an e-mail to your professor.

---

# Homework Questions:

## Working on loops and file interactions

We are seeking to develop a program that searches for one specific character in a given input, and counts how many times the character appears in this input. 

**Be careful here!
The text input and the character we are looking for are given as arguments of this function. 
So do not define any input(prompt) element.
Just give inputs as it is done in the test example below!**

To count how many times the character appears, you will have to define a count variable and add 1 each time the character is found. 

Give 2 different outputs for the user: 
- One concerns when the character could not be found into the input given
And you should use an output message like: 
        XX could not be found in 'gnrzbruo' 

- The other one gives the counting of the character
and you should use an output message like: 
        XX Has been found {nb} times in {input}


The program should only return the count. 
So 0 if you couldn't find the character in the input.

Do not forgot to use tests to see if your program runs well.
The print output does not count into the test processing system.


```python
#Method find / index 
text = 'I am in class today'
text.find('a')

```

Output:
```
2
```

<!-- grade_id: search, points: 0, grade: false, solution: true -->
```python
def search(text,character):
    '''This function searches for one character, 
    defined as 2nd argument, 
    in a given text / string of characters, 
    defined as 1 st argument, 
    and counts how many times it appears'''
    
    # The 1st argument text here represents 
    #     the input in which you are looking for 
    # The 2nd argument character represents 
    #     the character you are looking for
        
    #BEGIN SOLUTION
    err_msg = "Error: Your input is not transformable \
            into a string of characters"
    try: 
        text = str(text)
        character = str(character)
    except:
        print(err_msg)
        return None
        # In fact, this never happens with numbers (floats & integers)
        # It could happen with some specific objects like lists, dictionaries, ... 
    
    count = 0
    if character in text:
        for i in text:
            if i.lower() == character.lower():
                count += 1 #  count = count + 1
    else: 
        print(character+" could not be found in '{input}'".format(input=text))
        return count

    print(character+" Has been found {nb} times in '{input}'".format(nb=count,input=text))
    return count
    #END SOLUTION

    
```

```python
text= "Bonjour à tous, aujourd'hui nous donnons le cours en français, parce que tout le monde parle français. Ouuuaaiiiis!"
character = "e"

search(text,character)
```

Output:
```
e Has been found 7 times in 'Bonjour à tous, aujourd'hui nous donnons le cours en français, parce que tout le monde parle français. Ouuuaaiiiis!'
7
```

<!-- grade_id: cell-824190667948a132, points: 2, grade: true, solution: false -->
```python
import mock
from io import StringIO
import sys

text = "I've got 3 cats at home. And all 3 of them are lazy."
character = 3
assert search(text,character) == 2
assert search(text,'a') == 6

```

Output:
```
3 Has been found 2 times in 'I've got 3 cats at home. And all 3 of them are lazy.'
a Has been found 6 times in 'I've got 3 cats at home. And all 3 of them are lazy.'
```

<!-- grade_id: cell-42b99048c03d782a, points: 2, grade: true, solution: false -->
```python
import mock
from io import StringIO
import sys

text = "I've got 3 cats at home. And all 3 of them are lazy."
character = 2
assert search(text,character) == 0

```

Output:
```
2 could not be found in 'I've got 3 cats at home. And all 3 of them are lazy.'
```

---


## Exercises with strings and loops

Create a function here that takes 2 arguments (text & character).
The 1st one represents a text input given by the user.
The 2nd argument represents one character (only one letter here) to be found in the text.

This function takes 2 arguments. The 1st argument is a text and the 2nd argument is a character we look for inside the text argument. The function calculates some statistics (counting the number of words, the average length of words used and counts how many times the argument character appears)
    
If the input is not a text,
print a "Warning" message
and transform the input as string of characters.
This part is already coded here 
But use the flag after to change the output prints.

If the input is a text, 
Print out 
- The number of words 
- The average length of the words 
- How many times the character input appears in text.

**The function should return the 3 results from calculations.**

If the text is empty, return None and print out an error message.


You can use here the split() method applied to text. 

This method will split your text into multiple pieces. The delimiter should be defined as argument of the split() function. 
See an example below

For more details: help(str.split)


```python
help(str.split)
```

Output:
```
Help on method_descriptor:

split(self, /, sep=None, maxsplit=-1)
    Return a list of the words in the string, using sep as the delimiter string.
    
    sep
      The delimiter according which to split the string.
      None (the default value) means split according to any whitespace,
      and discard empty strings from the result.
    maxsplit
      Maximum number of splits to do.
      -1 (the default value) means no limit.
```

```python
text = 'This is a text'
text.split("i")

```

Output:
```
['Th', 's ', 's a text']
```

<!-- grade_id: statistics, points: 0, grade: false, solution: true -->
```python
def statistics(text,character):
    '''This function takes 2 arguments. The 1st argument is a text and the 2nd argument is a character 
    we look for inside the text argument. The function calculates some statistics 
    (counting the number of words, the average length of words used
    and counts how many times the argument character appears)'''
    
    warning_msg = "Warning: Your input is not a text. So we transformed it into one"
    
    # Giving here the warning message if the arguments are 
    if type(text)!= str and type(character)!=str:
        text = str(text)
        character = str(character)
        print(warning_msg)

    # Here we split the input into multiple parts when separated by a space
    text_list = text.split(" ")
    print(text_list)
    
    #BEGIN SOLUTION
    count_words = 0
    count_character =0
    sumlen = 0
    avglen = 0
    
    for word in text_list:
        count_words+=1   # count_words=count_words+1

        sumlen += len(word) # Here you can delete the punctuation not to count it if you want
        if character in word :
            #
            #####
            # Be sure here to count not only a character but a string 
            ####
            #
            count_character +=1
    if text.strip() == '':
        print('Error: Your input is empty')
        return None
    else: 
        avglen = sumlen / count_words
        print("The text counts {nbwords} words.".format(nbwords=count_words))
        print("The average length of the words is "+str(avglen)+" letters")
        print("{char} appears {countc} times in your input".format(char=character,countc=count_character))
        return count_words, avglen, count_character
    #END SOLUTION 


```

<!-- grade_id: cell-7a620da29566e996, points: 2, grade: true, solution: false -->
```python
lyrics = "I see trees of green Red roses too I see them bloom For me and you And I think to myself\
 What a wonderful world I see skies of blue And clouds of white The bright blessed day\
 The dark sacred night And I think to myself What a wonderful world The colors of the rainbow\
 So pretty in the sky Are also on the faces Of people going by I see friends shaking hands\
 Saying, How do you do? They are really saying I love you I hear babies cry\
 I watch them grow They will learn much more Than I will never know And I think to myself\
 What a wonderful world Yes, I think to myself What a wonderful world Oh yeah"

statistics(lyrics, 'is')

```

Output:
```
['I', 'see', 'trees', 'of', 'green', 'Red', 'roses', 'too', 'I', 'see', 'them', 'bloom', 'For', 'me', 'and', 'you', 'And', 'I', 'think', 'to', 'myself', 'What', 'a', 'wonderful', 'world', 'I', 'see', 'skies', 'of', 'blue', 'And', 'clouds', 'of', 'white', 'The', 'bright', 'blessed', 'day', 'The', 'dark', 'sacred', 'night', 'And', 'I', 'think', 'to', 'myself', 'What', 'a', 'wonderful', 'world', 'The', 'colors', 'of', 'the', 'rainbow', 'So', 'pretty', 'in', 'the', 'sky', 'Are', 'also', 'on', 'the', 'faces', 'Of', 'people', 'going', 'by', 'I', 'see', 'friends', 'shaking', 'hands', 'Saying,', 'How', 'do', 'you', 'do?', 'They', 'are', 'really', 'saying', 'I', 'love', 'you', 'I', 'hear', 'babies', 'cry', 'I', 'watch', 'them', 'grow', 'They', 'will', 'learn', 'much', 'more', 'Than', 'I', 'will', 'never', 'know', 'And', 'I', 'think', 'to', 'myself', 'What', 'a', 'wonderful', 'world', 'Yes,', 'I', 'think', 'to', 'myself', 'What', 'a', 'wonderful', 'world', 'Oh', 'yeah']
Error: Your input is empt
```

<!-- grade_id: cell-752bdae33ca2530a, points: 2, grade: true, solution: false -->
```python
# Testing
import mock
from io import StringIO
import sys

with mock.patch('sys.stdout', new_callable=StringIO):
    assert statistics(' ','i') is None

```

Output:
```
AssertionError                            Traceback (most recent call last)
/tmp/ipykernel_586944/123514496.py in <module>
      5 
      6 with mock.patch('sys.stdout', new_callable=StringIO):
----> 7     assert statistics(' ','i') is None

AssertionError:
```

<!-- grade_id: cell-a690ce56ba34b23a, points: 2, grade: true, solution: false -->
```python
# Testing
import mock
from io import StringIO
import sys

lyrics = "I see trees of green Red roses too I see them bloom For me and you And I think to myself\
 What a wonderful world I see skies of blue And clouds of white The bright blessed day\
 The dark sacred night And I think to myself What a wonderful world The colors of the rainbow\
 So pretty in the sky Are also on the faces Of people going by I see friends shaking hands\
 Saying, How do you do? They are really saying I love you I hear babies cry\
 I watch them grow They will learn much more Than I will never know And I think to myself\
 What a wonderful world Yes, I think to myself What a wonderful world Oh yeah"

with mock.patch('sys.stdout', new_callable=StringIO):
    assert statistics(lyrics,'is') == (125,3.792,0)

```

---


## Exercises with "Files"

This is a big important chapter.  All data files have to be loaded and parsed into Python before you can do analysis on them.  And you can also write out your results in files after doing this chapter.


**Question** Save the file mbox-short.txt from the chapter to your computer. Be sure you understand where it is located so you can read it in. If it is in the same directory as this notebook, you don't need to give a path to the file.  A path is the pwd info + filename from the Unix practice you did in class.  For instance:

```
myfile = "/Users/loeillet/Documents/MK34P/mbox-short.txt"
```


Create a function that reads in the file and looks for lines with "http" and "@" in them. (Hint: "http" can appear inside a longer string like "https".) Keep a count of how many URLs and how many email addresses appear. Print out the first URL in each line, but nothing else in that line. This means you need to isolate the URL from the rest of the line (you can use space as the delimiter.)

At the end, return:
- the total count of lines containing http, 
- the total count of lines containing @ in the file.

(You can use "return variable1, variable2" to return 2 variables from a function. 
It could be done with more the number of variables you want of course.)

** Note: the preferred way to do the file open and read is like this, not as shown in the book.  Please use this model in your code: **




```python
## Example of how to read in the file and do things with it:

# you need to give the full path to the file you are opening in the string, or a path relative to where your notebook
# A relative path means a folder + filename inside the current folder where the notebook is running.
# Here mine is in a folder called text-files.
file= 'mbox-short.txt'
with open(file, 'r') as handle:
    for line in handle:
        # do stuff here with each line.  Try uncommenting the print and removing the 'continue.  
        # Think about removing the \n at the end of the line.
        # print(line)
        continue

```

<!-- grade_id: count_urls_emails, points: 0, grade: false, solution: true -->
```python
# Other version 
def count_urls_emails(filename):
    """
    Pass in the full path to the file you are searching for URLs and Emails.
    Return the total of lines with http and @ found.
    """

    #BEGIN SOLUTION
    total_url_lines = 0
    total_email_lines = 0
    with open(filename, 'r') as handle:
        for line in handle:
            if 'http' in line: 
                total_url_lines +=1
                index = line.find('http') # This only to isolate the first url in the line
                space_index = line[index:].find(" ")
                print(line[index:index+space_index])

            
            if '@' in line: 
                total_email_lines +=1        
            #total_email_lines += '@' in line # List-comprehension way to code
            
    return total_url_lines, total_email_lines
    #END SOLUTION
```

<!-- grade_id: cell-a5c2aabad070d2ef, points: 2, grade: true, solution: false -->
```python
# test your code on your function -- put the path to the file in the argument.

filepath = 'mbox-short.txt'
count_urls_emails(filepath)

```

Output:
```
http://source.sakaiproject.org/viewsvn/?view=rev&rev=39772

https://collab.sakaiproject.org/portal) from the Source site.

http://source.sakaiproject.org/viewsvn/?view=rev&rev=39771

https://collab.sakaiproject.org/portal) from the Source site.

http://source.sakaiproject.org/viewsvn/?view=rev&rev=39770

https://source.sakaiproject.org/svn/site-manage/trunk/

https://collab.sakaiproject.org/portal) from the Source site.

http://source.sakaiproject.org/viewsvn/?view=rev&rev=39769

https://collab.sakaiproject.org/portal) from the Source site.

http://source.sakaiproject.org/viewsvn/?view=rev&rev=39766

https://collab.sakaiproject.org/portal) from the Source site.

http://source.sakaiproject.org/viewsvn/?view=rev&rev=39765

https://collab.sakaiproject.org/portal) from the Source site.

http://source.sakaiproject.org/viewsvn/?view=rev&rev=39764

https://source.sakaiproject.org/svn/msgcntr/trunk

http://jira.sakaiproject.org/jira/browse/SAK-12488

https://collab.sakaiproject.org/portal) fro
```

<!-- grade_id: cell-dd92fffc8a5359d5, points: 2, grade: true, solution: false -->
```python
# Tests for your code. Manually grading the url printing.  Testing the output count is correct.

assert count_urls_emails('data-files/mbox-short.txt') == (75,336)

```

Output:
```
https://collab.sakaiproject.org/portal)

https://collab.sakaiproject.org/portal)


https://collab.sakaiproject.org/portal)

https://collab.sakaiproject.org/portal)

https://collab.sakaiproject.org/portal)

https://collab.sakaiproject.org/portal)


http://jira.sakaiproject.org/jira/browse/SAK-12488
https://collab.sakaiproject.org/portal)


http://jira.sakaiproject.org/jira/browse/SAK-12484
https://collab.sakaiproject.org/portal)

http://bugs.sakaiproject.org/jira/browse/SAK-12596
https://collab.sakaiproject.org/portal)

http://bugs.sakaiproject.org/jira/browse/SAK-12595
https://collab.sakaiproject.org/portal)

https://collab.sakaiproject.org/portal)

http://bugs.sakaiproject.org/jira/browse/SAK-12592
https://collab.sakaiproject.org/portal)

http://bugs.sakaiproject.org/jira/browse/SAK-12175
https://collab.sakaiproject.org/portal)

https://collab.sakaiproject.org/portal)

https://collab.sakaiproject.org/portal)

https://collab.sakaiproject.org/portal)


https://source.sakaiproject.org/sv
```

```python
a = [1,2,3]
b = [4,5,6]
print(a+b)
```

Output:
```
[1, 2, 3, 4, 5, 6]
```

```python
a.append(b)
```

```python
print(b)
```

Output:
```
[4, 5, 6]
```

```python
print(b.extend(a))
```

Output:
```
None
```

```python
print(b)
```

Output:
```
[4, 5, 6, 1, 2, 3]
```

```python
c = b.pop(1)
print(c)
```

Output:
```
6
```

```python

```

```python
print(b)
```

Output:
```
[4, 1, 1, 2, 3]
```

```python
b.reverse()
```

```python
b.sort()
```

```python
a = ["a","b","c"]
b = ['d','e','f']
a.extend(b)
print(a.index("d"))
print(len(a))
a.pop(3)
a.sort()
a.reverse()
print(a)
```

Output:
```
3
6
['f', 'e', 'c', 'b', 'a']
```

```python
a = [1,2,3,4,5,6,7,8]
print(a[1::3])
```

Output:
```
[2, 5, 8]
```

```python
a = "this is a test"
b = [i for i in a]
```

```python
print(b)
```

Output:
```
['t', 'h', 'i', 's', ' ', 'i', 's', ' ', 'a', ' ', 't', 'e', 's', 't']
```

```python
a = "   "
a.strip()
```

Output:
```
''
```

---

```python
a = "this is a text"
b = a.split()
```

```python
print(b)
```

Output:
```
['this', 'is', 'a', 'text']
```

```python
" ".join(b)
```

Output:
```
'this is a text'
```

```python
a = [1,2,3]
b = [i+3 for i in a]
print(b)
```

Output:
```
[4, 5, 6]
```

```python

```