You should have read all the readings and content from the lectures given on Brightspace.


In order to make it easy to grade your works, here are the notes you should read and follow carrefully.  We will take off points for not following them. We will grade based on assuming you have read and understood them.
    
1. **Using the Notebook to Show Your Work**: You must learn to write code in the notebook... It is a core tool for data science and will make it easier to develop and document your work if you become good at using it. Writing your code in another tool and pasting it into the notebook will probably not work well: Some of you forgot to include code elsewhere, or the spacing became messed up, or it didn't run as copied.  You must be sure your code cells execute because we will test them.  So learn how to run code in the notebook cells to double check your work.
2. **Read Directions Carefully**:  The instructions for your code are very important. If you don't follow the requirements, your application won't do as requested. Making it work correctly is part of learning to program.  If we worded something unclearly, send an e-mail to your professor.


---


# Homework Questions:

## Working on lists and file interactions

Note: 
This focus doesn't mean the concepts we use are not cumulative.

You will see that all elements you practiced before will be useful. 

#### Some practice with lists to start

1. Create two lists: 

    - list "a" with integers from 0 to 9, 
    - list "b" computing x^2 pour x de 0 à 10 (inclus) avec un pas de 2


<!-- grade_id: cell-c394be06dcca0002, points: 1, grade: true, solution: true -->
```python
a = [x for x in range(10)]
b = [x**2 for x in range(0,11,2)]
a, b
```

Output:
```
([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], [0, 4, 16, 36, 64, 100])
```

<!-- grade_id: cell-7f5b1213545d213e, points: 1, grade: true, solution: false -->
```python
import mock
assert type(a)==list
assert type(b)==list
assert len(a)== 10
assert len(b)==6
assert a == [0,1,2,3,4,5,6,7,8,9]
assert b == [0, 4, 16, 36, 64, 100]

```

2. Compute some statistics on those 2 lists using the following methods: 

   sum / mean / median

**Note**: Please check the python documentation for more details on how to use those ones, the associated syntax, the arguments, etc



```python
# Execute this cell before computing statistics 
from statistics import mean,median

```

<!-- grade_id: cell-eb4927f736f5c0a8, points: 0, grade: false, solution: true -->
```python
# Compute results here

```

<!-- grade_id: cell-51e46bd8dc0c5658, points: 1, grade: true, solution: false -->
```python
#Some tests
assert type(median(a))==float
assert type(median(b))==float

```

<!-- grade_id: cell-b25b7ebfbd99fffa, points: 2, grade: true, solution: false -->
```python
assert median(a)==4.5
assert median(b)==26.0
assert sum(a)==45
assert sum(b)==220
assert mean(a)==4.5
assert round(mean(b),2)==36.67

```

-----


### Working with a NBA dataset 

We will focus this time on a dataset containing data about NBA players since 90's. 

You will read et extract data from the file and manipulate them using lists. 

We are seeking to explore data containing in this NBA dataset. We will focus on some specific columns and compute some statistics using list methods as seen during the course (and some you can explore by yourselves in the documentation). 

--

**Nota: 
Using the documentation (and info you can find on the internet - like on stackoverflow) is truely part of the development of your programming skills (and not only with Python). 

This is why, sometimes you will find some questions using some methods that we didn't discuss during class, in order to push you to explore those by yourselves.

--


In this first question, 
Create a function that extract specific data from the csv file. 
The function takes 2 arguments: 
    - one to specify the path of your file
    - the other to indicate which columns you want to extract data from

The program should return a list of elements (the ones which are extracted from the csv file).


Once this function will be running, 
compute some statistics about data from the file 
(See below)

--

**Tip:**
Be careful when extracting data using split method when creating list from lines of the file. 
Check the type of the elements you extracted from the file. 

**Do not forgot to use tests to see if your program runs well.
The print output does not count into the test processing system.**

**Of course, do not forgot to test the path.**

If this one doesn't exist, print an error message and quit the function.


You can use here the split() method applied to text. 

This method will split your text into multiple pieces. The delimiter should be defined as argument of the split() function. 
See an example below

For more details: help(str.split)


```python
help(str.split)

```

Output:
```
Help on method_descriptor:

split(self, /, sep=None, maxsplit=-1)
    Return a list of the words in the string, using sep as the delimiter string.
    
    sep
      The delimiter according which to split the string.
      None (the default value) means split according to any whitespace,
      and discard empty strings from the result.
    maxsplit
      Maximum number of splits to do.
      -1 (the default value) means no limit.
```

```python
# One example of using split method on string object
text = 'This is a text'
text.split("i")

```

Output:
```
['Th', 's ', 's a text']
```

Another option here can use a specific method to read csv file. 

This is : **csv.reader**
from the csv library

To use this, you will need to import the library on top of your cell, using: 

    import csv


```python
import csv
help(csv.reader)

```

Output:
```
Help on built-in function reader in module _csv:

reader(...)
    csv_reader = reader(iterable [, dialect='excel']
                            [optional keyword args])
        for row in csv_reader:
            process(row)
    
    The "iterable" argument can be any object that returns a line
    of input for each iteration, such as a file object or a list.  The
    optional "dialect" parameter is discussed below.  The function
    also accepts optional keyword arguments which override settings
    provided by the dialect.
    
    The returned object is an iterator.  Each iteration returns a row
    of the CSV file (which can span multiple input lines).
```

<!-- grade_id: cell-b08adcbee18d15bc, points: 3, grade: true, solution: true -->
```python
# Write here your function

def extract_data(file,index):
    '''This function takes a csv file as argument.
    It extracts data from the NBA dataset file.
    The argument index specifies which column you want to extract. 
    The function returns a list of data containing all data 
    extracted from the csv file.
    '''
        
    #BEGIN SOLUTION
    # First version to read csv file (not the most efficient)

    ## Initialisation:
    datalist = []
    flag=0 # Defining a flag to ignore the first line in the file 
    # that corresponds to column titles
    
    err_msg = "Error: Your file doesn't exist nor a csv format"
    
    #try: 
    with open(file,'r', encoding='utf-8') as handle:
        for i in handle: 
            if flag ==0 : 
                flag +=1
                continue
            line_list = i.split(',')
            datalist.append(line_list[index])
                
        # As list comprehension: 
            #datalist = [i.split(',') for i in handle if flag==0: continue]
            #flag +=1
        
#    except:
#        print(err_msg)
#        return None
    
    if index not in (0,1,2,6,7,8,9,10,11,21):
        datalist = [float(i) for i in datalist]
    
    return datalist
    #END SOLUTION

```

```python
# Define for which columns you want to extract data from the file
index = 4

# Specify the path of your NBA dataset file 
file = 'all_seasons.csv'

extract_data(file,index)

```

Output:
```
8.18457264957265
```

<!-- grade_id: cell-7a4ef495d5bc5cd4, points: 0, grade: false, solution: true -->
```python
# Compute here some statistics about the height and weight of the NBA players over the years: 
#    - average 
#    - median

# Comment the process that drives you to that result
#BEGIN SOLUTION 

#END SOLUTION

```

<!-- grade_id: cell-0c71068b874fa866, points: 1, grade: true, solution: false -->
```python
# Some test for your new function
import mock
from io import StringIO
import sys

index=2
assert type(extract_data(file,index)) == list

```

<!-- grade_id: cell-8394990dfbae81b5, points: 2, grade: true, solution: false -->
```python
# Some other tests about data statistics from the file

index = 4
assert round(mean(extract_data(file,index)),2) == 200.73
index = 5
assert round(mean(extract_data(file,index)),2) == 100.53

```

---

## Second step

Get inspired by what you've done on previous step! 

You saw that with that kind of function you need to only focus on one specific column. It is limited indeed. 

So now, you will try to take benefit from lists to extract data from few columns from the csv file and pack it all in one main list. The idea is to create a list of lists in which each element of the main list contains a list of data from each selected column. 

So let's a new function that takes two arguments: 
    - First argument: still the path of the csv file 
    - Second argument: the indexes that you want to extract. This argument should be a list of the indexes of selected columns 

**Tip:**
You can of course use the function you created above.


<!-- grade_id: cell-b61e501e4cdf89a7, points: 2, grade: true, solution: true -->
```python
def extract_fewdata(file,index):

    error_msg = 'The index argument is not a list'
    
    # BEGIN SOLUTION
    if type(index)!= list:
        print(error_msg)
        return None
    
    data = []
    for i in index: 
        data.append(extract_data(file,i))
        
    return data
    # END SOLUTION

```

```python
index = [4,5]
file = 'all_seasons.csv'
extract_fewdata(file,index)

```

Output:
```
[[213.36,
  210.82,
  208.28,
  210.82,
  205.74,
  198.12,
  208.28,
  208.28,
  182.88,
  190.5,
  205.74,
  200.66,
  205.74,
  208.28,
  195.58,
  210.82,
  190.5,
  203.2,
  200.66,
  185.42,
  195.58,
  218.44,
  205.74,
  200.66,
  215.9,
  190.5,
  200.66,
  187.96,
  205.74,
  213.36,
  210.82,
  208.28,
  205.74,
  205.74,
  213.36,
  185.42,
  198.12,
  195.58,
  195.58,
  205.74,
  160.02,
  187.96,
  203.2,
  195.58,
  195.58,
  208.28,
  203.2,
  198.12,
  182.88,
  200.66,
  195.58,
  205.74,
  185.42,
  185.42,
  210.82,
  200.66,
  203.2,
  205.74,
  205.74,
  198.12,
  195.58,
  200.66,
  205.74,
  185.42,
  195.58,
  195.58,
  187.96,
  210.82,
  200.66,
  208.28,
  213.36,
  190.5,
  203.2,
  195.58,
  195.58,
  208.28,
  215.9,
  198.12,
  198.12,
  210.82,
  203.2,
  195.58,
  185.42,
  185.42,
  200.66,
  187.96,
  185.42,
  195.58,
  195.58,
  203.2,
  205.74,
  200.66,
  200.66,
  208.28,
  208.28,
  203.2,
  200.66,
  190.5,
  213.36,
  213.36,
  185.42,
  210
```

<!-- grade_id: cell-27afdaf764719668, points: 2, grade: true, solution: false -->
```python
index = [4,5]
assert len(extract_fewdata(file,index)) == 2
assert len(extract_fewdata(file,index)[0]) == 11700

```

---


## Third step

Well done ! 
Let's get inspired again by what you've done on previous steps! 

So now, you will try to generalize. Still, the idea is to create a list of lists in which each element of the main list contains a list of data from each column this time. 

That means you will have to: 

Create a function that create a list of lists. 
The function takes 1 argument only which corresponds to the path of the csv file.  


**The function should return the main list.**

You can manipulate the list and containing data afterwards. 

If the file is empty, return None and print out an "warning" message.


<!-- grade_id: cell-1d2224a69e2996d9, points: 2, grade: true, solution: true -->
```python
def extract_fulldata(file):
    '''This function takes 1 argument which corresponds 
    to the csv file path. The function creates one main list 
    containing all extracted data, separated in multiple lists. 
    '''
    
    warning_msg = "Warning: The given csv file is empty!"
    
    #BEGIN SOLUTION
    ## Initialisation:
    with open(file,'r', encoding='utf-8') as handle:
        for i in handle:
            length = len(i.split(','))
            break
        index = list(range(length-1))
    
    return extract_fewdata(file,index)
    #END SOLUTION

```

```python
# Specify the path of your NBA dataset file 
file = 'all_seasons.csv'

mean(extract_fulldata(file)[12])
```

Output:
```
8.18457264957265
```

<!-- grade_id: cell-da803813eea90ca3, points: 1, grade: true, solution: false -->
```python
# Testing
import mock
assert type(extract_fulldata(file))==list

```

<!-- grade_id: cell-47d571fd9ee0e2f5, points: 2, grade: true, solution: false -->
```python
# More testing

assert round(mean(extract_fulldata(file)[4]),2)==200.73
assert round(mean(extract_fulldata(file)[12]),2)==8.18
```

---