Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All). Check your output to make sure it all looks as you expected.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "YOUR NAME HERE"
```

---

You should have read all the readings and content from the lectures given on Brightspace.


In order to make it easy to grade your works, here are the notes you should read and follow carrefully.  We will take off points for not following them. We will grade based on assuming you have read and understood them.
    
1. **Using the Notebook to Show Your Work**: You must learn to write code in the notebook... It is a core tool for data science and will make it easier to develop and document your work if you become good at using it. Writing your code in another tool and pasting it into the notebook will probably not work well: Some of you forgot to include code elsewhere, or the spacing became messed up, or it didn't run as copied.  You must be sure your code cells execute because we will test them.  So learn how to run code in the notebook cells to double check your work.
2. **Read Directions Carefully**:  The instructions for your code are very important. If you don't follow the requirements, your application won't do as requested. Making it work correctly is part of learning to program.  If we worded something unclearly, send an e-mail to your professor.

---

# Homework Questions:

## Working on dictionaries and file interactions
(and previous content ;) )

Note: 
This focus doesn't mean the concepts we use are not cumulative.

You will see that all elements you practiced before will be useful. 

## First exercise: 

### Working with a NBA dataset 

Back to our NBA dataset.

You will read et extract data from the file and manipulate them but this time using dictionaries.

The exercise here consists in creating a function that :
- take a file as argument
- build up a dictionary containing all data from the file

    The keys of the dictionary should be the title of columns. 

    And value for each key should be a list of data from the same column.


**Note: 
The main function should be call dict_data(file) !**

**Tip: 
I advise you to consider multiple functions (as we did in the last assignment). But it can also be done with one (with more steps and less "pythonish")**

<!-- grade_id: cell-23a803a5d9214c0b, points: 4, grade: true, solution: true -->
```python
import csv 

err_msg = "Error: the file has not be found"
def extract_data(file,index):
    datalist = []
    try:
        with open (file, 'r') as handle:
            csv_reader = csv.reader(handle)
            for row in csv_reader:
                datalist.append(row[index])
    except:
        print(err_msg)
        return 
    numbers_list = []
    for i in datalist[1:]:
        try:
            numbers_list.append(float(i))
        except:
            numbers_list.append(i)
    return numbers_list


def dict_data(file):
    '''This function takes a csv file as argument.
    It extracts data from the NBA dataset file,
    And compiles all of them into a dictionary.
    The keys of this one corresponds to column titles.
    '''
        
    ## Initialisation:
    data = {}
    err_msg = "Error: Your file doesn't exist nor a csv format"
    
    try:
        with open(file,'r', encoding='utf-8') as handle:
            content = csv.reader(handle)
            # Reading and saving only the first row 
            # (=column titles)
            titles = next(content)
#             for row in content:
#                 titles = row
#                 break
    except: 
        print(err_msg)
        return None
        
    # Exploring each column and extracting data from those
    # to put them into a structure dictionary
    index = list(range(len(titles)-1))    
    for i in index: 
        data[titles[i]] = extract_data(file,i)
                   
    return data
```

```python

```

```python
file = 'all_seasons.csv'
dict_data(file)

# You can now access data from specific columns 
# by giving into [] the name of the selected column.
```

Output:
```
{'Unnamed: 0': [0.0,
  1.0,
  2.0,
  3.0,
  4.0,
  5.0,
  6.0,
  7.0,
  8.0,
  9.0,
  10.0,
  11.0,
  12.0,
  13.0,
  14.0,
  15.0,
  16.0,
  17.0,
  18.0,
  19.0,
  20.0,
  21.0,
  22.0,
  23.0,
  24.0,
  25.0,
  26.0,
  27.0,
  28.0,
  29.0,
  30.0,
  31.0,
  32.0,
  33.0,
  34.0,
  35.0,
  36.0,
  37.0,
  38.0,
  39.0,
  40.0,
  41.0,
  42.0,
  43.0,
  44.0,
  45.0,
  46.0,
  47.0,
  48.0,
  49.0,
  50.0,
  51.0,
  52.0,
  53.0,
  54.0,
  55.0,
  56.0,
  57.0,
  58.0,
  59.0,
  60.0,
  61.0,
  62.0,
  63.0,
  64.0,
  65.0,
  66.0,
  67.0,
  68.0,
  69.0,
  70.0,
  71.0,
  72.0,
  73.0,
  74.0,
  75.0,
  76.0,
  77.0,
  78.0,
  79.0,
  80.0,
  81.0,
  82.0,
  83.0,
  84.0,
  85.0,
  86.0,
  87.0,
  88.0,
  89.0,
  90.0,
  91.0,
  92.0,
  93.0,
  94.0,
  95.0,
  96.0,
  97.0,
  98.0,
  99.0,
  100.0,
  101.0,
  102.0,
  103.0,
  104.0,
  105.0,
  106.0,
  107.0,
  108.0,
  109.0,
  110.0,
  111.0,
  112.0,
  113.0,
  114.0,
  115.0,
  116.0,
  117.0,
  118.0,
  119.0,
  120.0,
  121.0
```

<!-- grade_id: cell-596ebfefec494a60, points: 2, grade: true, solution: false -->
```python
import mock 
assert type(dict_data(file)) == dict
assert len(dict_data(file)) == 21
assert len(dict_data(file)['pts']) == 11700
```

2. Compute some statistics on those 2 lists using the following methods: 

   sum / mean / median

**Note**: You can compare the results you get to the ones you had in the previous assignement using lists.

```python
# Execute this cell before computing statistics 
from statistics import mean,median
```

```python
# Compute results here
# Compute results here
a1=mean(dict_data(file)['player_height'])
b1=mean(dict_data(file)['player_weight'])
a2=median(dict_data(file)['player_height'])
b2=median(dict_data(file)['player_weight'])
c1=sum(dict_data(file)['player_height'])
c2=sum(dict_data(file)['player_weight'])
print(a1)
print(b1)
```

Output:
```
200.72850085470085
100.52679086153846
```

<!-- grade_id: cell-d8b45844149d8ea6, points: 2, grade: true, solution: false -->
```python
assert round(mean(dict_data(file)['age']),2) == 27.13
assert round(mean(dict_data(file)['pts']),2) == 8.17
assert round(mean(dict_data(file)['reb']),2) == 3.56

```

-----

## Second exercise

### Working with a space mission dataset 

This is a very specific function to be programmed, but it will allow you to train on using loops, lists and dictionary.

We want to create a function that takes one argument which is a CSV file (here it only works with SpaceMission_data.csv file). 

The function explores the file and returns a dictionary in which it counts :
- The number of mission for each company (company will be keys of the dictionary)
- The number of successful missions 
- The number of failed missions

Each value element of the dictionary will be a list of the 3 counts described above. 
It will look like this: 
{company : [total number of missions, number of successful missions, number of failed missions]}

Explore the file first, and find a solution to identify the successful and failed missions identification. 

You will see that it is not only a question about a "," delimiter. 
So in order to handle easily the reading of the format, i strongly suggest you to use this method: 
Just after opening the file, 
use "content = csv.reader(handle)"

And you will be able to loop into content.

And of course return a dictionary! 

<!-- grade_id: search, points: 4, grade: true, solution: true -->
```python
import csv
def space_statistics(filename):
    '''This function takes one argument which should be a CSV file. 
    It specifically explores the space_corrected.csv file. 
    The function returns a dictionary in which it counts the number of mission for each company in charge
    and also counts the number of successful and failed missions'''

    space = dict()
    err_Msg = 'Error: Check the file path or content' 
 
    #BEGIN SOLUTION 
    try: 
        with open(filename,'r') as handle:
            content = csv.reader(handle)
            next(content)
                        
            for line in content: 
                                
                # Extracting the company name
                company = line[2]
                # Checking if the company name has already an entry in the dictionary
                if company not in space: 
                    # Counting if it is a success or kind of a failure (partial, or other)
                    if 'failure' in line[8].lower():
                        space[company]=[1,0,1]
                    elif 'success' in line[8].lower():
                        space[company]=[1,1,0]
                        
                else: 
                    if 'failure' in line[8].lower():
                        success_list = space[company]
                        success_list[0] += 1
                        success_list[2] += 1
                        space[company]= success_list
                    elif 'success' in line[8].lower():
                        success_list = space[company]
                        success_list[0] += 1
                        success_list[1] += 1
                        space[company]= success_list
    except: 
        print(err_msg)
        return None

    return space

    #END SOLUTION 

help(space_statistics)
```

Output:
```
Help on function space_statistics in module __main__:

space_statistics(filename)
    This function takes one argument which should be a CSV file. 
    It specifically explores the space_corrected.csv file. 
    The function returns a dictionary in which it counts the number of mission for each company in charge
    and also counts the number of successful and failed missions
```

```python
# Specify here the path of your SpaceMission_data file
file = 'SpaceMission_data.csv'
space_statistics(file)
```

Output:
```
{'SpaceX': [100, 94, 6],
 'CASC': [251, 231, 20],
 'Roscosmos': [55, 51, 4],
 'ULA': [140, 139, 1],
 'JAXA': [7, 6, 1],
 'Northrop': [83, 74, 9],
 'ExPace': [10, 9, 1],
 'IAI': [11, 9, 2],
 'Rocket Lab': [13, 11, 2],
 'Virgin Orbit': [1, 0, 1],
 'VKS RF': [201, 188, 13],
 'MHI': [84, 80, 4],
 'IRGC': [1, 1, 0],
 'Arianespace': [279, 269, 10],
 'ISA': [13, 4, 9],
 'Blue Origin': [12, 12, 0],
 'ISRO': [76, 63, 13],
 'Exos': [4, 0, 4],
 'ILS': [46, 45, 1],
 'i-Space': [1, 1, 0],
 'OneSpace': [1, 0, 1],
 'Landspace': [1, 0, 1],
 'Eurockot': [13, 12, 1],
 'Land Launch': [7, 6, 1],
 'CASIC': [5, 3, 2],
 'KCST': [5, 2, 3],
 'Sandia': [1, 0, 1],
 'Kosmotras': [22, 21, 1],
 'Khrunichev': [1, 1, 0],
 'Sea Launch': [36, 33, 3],
 'KARI': [3, 1, 2],
 'ESA': [13, 9, 4],
 'NASA': [203, 186, 17],
 'Boeing': [136, 131, 5],
 'ISAS': [30, 26, 4],
 'SRC': [3, 2, 1],
 'MITT': [7, 6, 1],
 'Lockheed': [79, 74, 5],
 'AEB': [3, 0, 3],
 'Starsem': [1, 1, 0],
 'RVSN USSR': [1777, 1614, 163],
 'EER': [1, 0, 1],
 
```

<!-- grade_id: cell-48c5c08fcf433ee9, points: 2, grade: true, solution: false -->
```python
# Testing the output 
import mock
from io import StringIO
import sys

space = {'SpaceX': [100, 94, 6], 'CASC': [251, 231, 20], 'Roscosmos': [55, 51, 4], 'ULA': [140, 139, 1], 'JAXA': [7, 6, 1], 'Northrop': [83, 74, 9], 'ExPace': [10, 9, 1], 'IAI': [11, 9, 2], 'Rocket Lab': [13, 11, 2], 'Virgin Orbit': [1, 0, 1], 'VKS RF': [201, 188, 13], 'MHI': [84, 80, 4], 'IRGC': [1, 1, 0], 'Arianespace': [279, 269, 10], 'ISA': [13, 4, 9], 'Blue Origin': [12, 12, 0], 'ISRO': [76, 63, 13], 'Exos': [4, 0, 4], 'ILS': [46, 45, 1], 'i-Space': [1, 1, 0], 'OneSpace': [1, 0, 1], 'Landspace': [1, 0, 1], 'Eurockot': [13, 12, 1], 'Land Launch': [7, 6, 1], 'CASIC': [5, 3, 2], 'KCST': [5, 2, 3], 'Sandia': [1, 0, 1], 'Kosmotras': [22, 21, 1], 'Khrunichev': [1, 1, 0], 'Sea Launch': [36, 33, 3], 'KARI': [3, 1, 2], 'ESA': [13, 9, 4], 'NASA': [203, 186, 17], 'Boeing': [136, 131, 5], 'ISAS': [30, 26, 4], 'SRC': [3, 2, 1], 'MITT': [7, 6, 1], 'Lockheed': [79, 74, 5], 'AEB': [3, 0, 3], 'Starsem': [1, 1, 0], 'RVSN USSR': [1777, 1614, 163], 'EER': [1, 0, 1], 'General Dynamics': [251, 203, 48], 'Martin Marietta': [114, 100, 14], 'Yuzhmash': [2, 2, 0], 'Douglas': [1, 1, 0], 'ASI': [9, 9, 0], 'US Air Force': [161, 129, 32], 'CNES': [8, 6, 2], 'CECLES': [4, 1, 3], 'RAE': [2, 1, 1], 'UT': [5, 1, 4], 'OKB-586': [2, 2, 0], 'AMBA': [8, 4, 4], "Arm??e de l'Air": [4, 3, 1], 'US Navy': [17, 2, 15]}

with mock.patch('sys.stdout', new_callable=StringIO):
    assert space_statistics(file) == space

```

<!-- grade_id: cell-824190667948a132, points: 2, grade: true, solution: false -->
```python
# Some test to verify if all lines have been counted
import mock
from io import StringIO
import sys

with mock.patch('sys.stdout', new_callable=StringIO):
    assert sum([space_statistics(file)[keys][0] for keys in space]) == 4324

```

<!-- grade_id: cell-42b99048c03d782a, points: 2, grade: true, solution: false -->
```python
# Some test to verify one result from the output dictionary
import mock
from io import StringIO
import sys

with mock.patch('sys.stdout', new_callable=StringIO):
    assert space_statistics(file)['SpaceX'] == [100,94,6]

```

## Congrats! You have done some datascience with Python !
You can also have alook at PANDAS (https://pandas.pydata.org). 
This Python librairy is mostly used in data science to explore, and play with data (from files, servers, or else).

For your information, The **MK347 course, Python data analysis** is dedicated to that at emlyon. 

---