### Creation of a currency converter ###

Consider money you need to convert into a different currency. 
We will ask you to code a converter by asking the user to give an input value, and indicate the conversion he or she wants to make (from the initial currency to the final one).

The initial and final currencies will also be inputs from the user. 


#### First step####

Let's consider you as an user want to convert money from EUR to YEN (Japanese currency).
Please take the currency value to be 1 EUR = 8.09647 YEN (on the 13/09/2020 - 1:00 pm / CET+2)

Write a code which: 
- Asks from the user an input considered as value money
This one will be considered to be in EUR 

- Then calculates the value in YEN 
- And displays and returns the result 
Here use both "print" and "return" functions

**Do not forgot to include tests / error messages / ...
in order to guide the user when using your program !**

All your error messages must start with "Error: " 


<!-- grade_id: cell-1b2e96b62467116f, points: 3, grade: true, solution: true -->
```python
# Remember you can comment your code using #

def currency_converter():
    ''' Ask for the user input with a prompt and compute the currency conversion from EUR to YEN. Print the value
    and return it from the function as well.'''
    
    prompt = 'Enter a value in EUR to be converted to YEN:'
    error_message = 'Error: your input should be a positive number'
    
    ### BEGIN SOLUTION
    value = input(prompt)
    try :
        value=float(value) 
    except:
        print(error_message)
        return
    
    if value >= 0 :
        conversion = value * 8.09647
    else: 
        print(error_message)
        return
    ### END SOLUTION

    print("Your input is equal to {output} YEN".format(output=conversion))
    return conversion

```

To run your code in a cell, you should add a new cell and type:

```currency_converter()```

and execute that cell.  Provide the input and show it working below:

```python
# run this to show it running!
currency_converter()

```

Output:
```
Enter a value in EUR to be converted to YEN:-170
Error: your input should be a positive number
```

The following code checks that your function performs correctly.  It is only verifying the return value is the correct number.  Testing functions for their return values is a standard practice in *Test-Driven Development*.  Test cases help you figure out where your code has errors, if you have lots of functions calling each other.

<!-- grade_id: test_converter1, points: 1, grade: true, solution: false -->
```python
# This is a test of your code. Don't change this.

import mock
from io import StringIO
import sys

def test_converter():
    with mock.patch('builtins.input', return_value='4.5'):
        assert currency_converter() == 36.434115
    with mock.patch('builtins.input', return_value='200'):
        assert currency_converter() == 1619.294
            
test_converter()
```

Output:
```
Your input is equal to 36.434115 YEN
Your input is equal to 1619.294 YEN
```

<!-- grade_id: test_converter2, points: 2, grade: true, solution: false -->
```python
# testing more....

with mock.patch('builtins.input', return_value='cat'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        assert currency_converter() is None
        assert sys.stdout.getvalue() == "Error: your input should be a positive number\n"
```

<!-- grade_id: cell-c92e17942d63eb09, points: 1, grade: true, solution: false -->
```python
# testing even more....

with mock.patch('builtins.input', return_value='-3'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        assert currency_converter() is None
        assert sys.stdout.getvalue() == "Error: your input should be a positive number\n"
```

#### Second step####

Let's continue on the same topic.
Now consider multiple options of conversion from one currency to another one. 

This means by using what you've done preciously, add 1 input from the user to be: 
- the final currency for which you need to know the value
The initial currency of the value given will be stuck to EUR. 

For that purpose, you need to : 
- Store inputs into variables 
- Use inputs as elements for conditionals 
- Calculates conversion depending on the inputs
- Only use one output variable

For this exercise, you will consider the following currencies: YEN / USD / CAD / GBP / INR
The conversion rates are: 
- 1 EUR = 1.18468 USD
- 1 EUR = 0.925698 GBP 
- 1 EUR = 87.0200 INR
- 1 EUR = 1.56086 CAD

**When asking for inputs from the user, please stick to upper letters**
Please consider here inputs as arguments to the function currency_converter2()

<!-- grade_id: cell-acd5cd5cae046cf1, points: 2, grade: true, solution: true -->
```python
# Remember you can comment your code using #

def currency_converter2(value,currency):
    ''' Ask for the user input using the arguments of the function and compute the currency conversion from EUR to YEN. Print the value
    and return it from the function as well.'''

    error_message = 'Error: your input should be a positive number'
    error_message2 = 'Error: select one of the given currencies (YEN / USD / CAD / GBP / INR)'

    ### BEGIN SOLUTION
    try :
        value=float(value) 
    except:
        print(error_message)
        return
    
    if value >= 0 :
        if currency == 'YEN':
            conversion = value * 8.09647
        elif currency == 'USD':
            conversion = value * 1.18468
        elif currency == 'CAD': 
            conversion = value * 1.56086
        elif currency == 'GBP':
            conversion = value * 0.925698
        elif currency == 'INR': 
            conversion = value * 87.0200
        else : 
            print(error_message2)
            return            
    else: 
        print(error_message)
        return
    ### END SOLUTION

    print("Your input is equal to {output} ".format(output=conversion)+currency)
    print('Your input is equal to',conversion,currency)
    return conversion

```

```python
currency_converter2(2,'YEN')
```

Output:
```
Your input is equal to 16.19294 YEN
Your input is equal to 16.19294 YEN
16.19294
```

<!-- grade_id: cell-486bfab8a6c7eea7, points: 0, grade: false, solution: true -->
```python
input1 = 12
input2 = 'GBP'
currency_converter2(input1, input2)
```

Output:
```
Your input is equal to 11.108376 GBP
Your input is equal to 11.108376 GBP
11.108376
```

<!-- grade_id: cell-160387a81d89364a, points: 0, grade: false, solution: true -->
```python
# run this to show if it runs correctly!
# Give here content for the 2 variables value and currency
input1 = 30.45
input2 = 'INR'
currency_converter2(input1, input2)

```

Output:
```
Your input is equal to 2649.759 INR
Your input is equal to 2649.759 INR
2649.759
```

<!-- grade_id: test1_currency_converter2, points: 1, grade: true, solution: false -->
```python
# This is a test of your code. Don't change this.

import mock
from io import StringIO
import sys

assert currency_converter2(3,'GBP') == 2.777094
assert currency_converter2(200,'USD') == 236.93599999999998
```

Output:
```
Your input is equal to 2.777094 GBP
Your input is equal to 2.777094 GBP
Your input is equal to 236.93599999999998 USD
Your input is equal to 236.93599999999998 USD
```

<!-- grade_id: test2_currency_converter2, points: 2, grade: true, solution: false -->
```python
# Some more tests
assert currency_converter2('cat','GBP') == None
assert currency_converter2(3,'Other') == None

```

Output:
```
Error: your input should be a positive number
Error: select one of the given currencies (YEN / USD / CAD / GBP / INR)
```

<!-- grade_id: test3_currency_converter2, points: 1, grade: true, solution: false -->
```python
assert currency_converter2(0,'GBP') == 0

```

Output:
```
Your input is equal to 0.0 GBP
Your input is equal to 0.0 GBP
```

#### Third step####

During the last step, you will extend once again the program you've beeing programming. 

**Now consider not only one freezed input currency but multiple. 
This means the user need to specify in which currency the input is. 
He or she still has also to specify the output currency.** 

So here you need to: 
- Add one more input from the user specifying the input or initial currency.

Once again store this input into a variable and use it for calculating the output value. 

**Take time before coding to consider all the options of conversion, from one currency to another one. And try a way to calculate each conversion even if it is not given here in the instructions.** 

Concerning the conversion rates, use the ones indicate in 2nd step. 



<!-- grade_id: cell-e6ae9b5ba8031bea, points: 3, grade: true, solution: true -->
```python
# Remember you can comment your code using #
# You can see that here we switch from input asking to arguments giving directly to the function

def currency_converter3(value, currency1, currency2):
    ''' From a value given by the user the function computes the currency conversion from one currency to another one. Print the value
    and return it from the function as well.'''
    
    error_message = 'Error: your input should be a positive number'
    error_message2 = 'Error: select one of the given currency (EUR,YEN,USD,CAD,INR,GBP)'

    ### BEGIN SOLUTION
    try :
        value=float(value)
    except:
        print(error_message)
        return
    
    EURYEN = 8.09647
    EURUSD = 1.18468
    EURCAD = 1.56086
    EURGBP = 0.925698
    EURINR = 87.0200
    
    if value >= 0 :
        if currency1 in ('EUR','YEN','USD','CAD','INR','GBP') and currency2 in ('EUR','YEN','USD','CAD','INR','GBP'):    
            if currency1 == 'EUR':
                if currency2 == 'YEN':
                    conversion = value * EURYEN
                elif currency2 == 'USD':
                    conversion = value * EURUSD
                elif currency2 == 'CAD':
                    conversion = value * EURCAD
                elif currency2 == 'GBP':
                    conversion = value * EURGBP                    
                elif currency2 == 'INR':
                    conversion = value * EURINR
                elif currency2 == 'EUR':
                    conversion = value 
            if currency1 == 'YEN':
                if currency2 == 'EUR':
                    conversion = value * 1/EURYEN
                elif currency2 == 'USD':
                    conversion = value * EURUSD/EURYEN
                elif currency2 == 'CAD':
                    conversion = value * EURCAD/EURYEN
                elif currency2 == 'GBP':
                    conversion = value * EURGBP/EURYEN
                elif currency2 == 'INR':
                    conversion = value * EURINR/EURYEN
            if currency1 == 'USD':
                if currency2 == 'YEN':
                    conversion = value * EURYEN/EURUSD
                elif currency2 == 'EUR':
                    conversion = value * 1/(EURUSD)
                elif currency2 == 'CAD':
                    conversion = value * EURCAD/EURUSD
                elif currency2 == 'GBP':
                    conversion = value * EURGBP/EURUSD
                elif currency2 == 'INR':
                    conversion = value * EURINR/EURUSD
            if currency1 == 'CAD':
                if currency2 == 'YEN':
                    conversion = value * EURYEN/EURCAD
                elif currency2 == 'EUR':
                    conversion = value * 1/(EURCAD)
                elif currency2 == 'USD':
                    conversion = value * EURUSD/EURCAD
                elif currency2 == 'GBP':
                    conversion = value * EURGBP/EURCAD
                elif currency2 == 'INR':
                    conversion = value * EURINR/EURCAD
            if currency1 == 'GBP':
                if currency2 == 'YEN':
                    conversion = value * EURYEN/EURGBP
                elif currency2 == 'EUR':
                    conversion = value * 1/(EURGBP)
                elif currency2 == 'USD':
                    conversion = value * EURCAD/EURYEN
                elif currency2 == 'CAD':
                    conversion = value * EURCAD/EURYEN
                elif currency2 == 'INR':
                    conversion = value * EURINR/EURYEN
            if currency1 == 'INR':
                if currency2 == 'YEN':
                    conversion = value * EURYEN/EURINR
                elif currency2 == 'GBP':
                    conversion = value * EURGBP/EURINR
                elif currency2 == 'USD':
                    conversion = value * EURUSD/EURINR
                elif currency2 == 'CAD':
                    conversion = value * EURCAD/EURINR
                elif currency2 == 'EUR':
                    conversion = value * 1/(EURINR)
        else : 
            print(error_message2)
            return            
    else: 
        print(error_message)
        return
    ### END SOLUTION

    print(conversion)
    return conversion

```

```python

```

<!-- grade_id: cell-c788fe9305e7f4f7, points: 2, grade: true, solution: false -->
```python
# This is a test of your code. Don't change this.

import mock
from io import StringIO
import sys

assert currency_converter3(200,'INR','USD') == 2.722776373247529
assert currency_converter3(30,'CAD','GBP') == 17.792076163140834

```

Output:
```
2.722776373247529
17.792076163140834
```

<!-- grade_id: cell-16120a64ef493730, points: 2, grade: true, solution: false -->
```python
# Some more tests
assert currency_converter3(3,'cat','GBP') == None
assert currency_converter3(3,'EUR','Other') == None

```

Output:
```
Error: select one of the given currency (EUR,YEN,USD,CAD,INR,GBP)
Error: select one of the given currency (EUR,YEN,USD,CAD,INR,GBP)
```